# CSS 201.5 - CSS Bootcamp

## Linear Algebra

### Exam: Computational Linear Algebra

In this exam, you'll get hands-on practice with the following **concepts**:

- Classes
- Functions
- Matrix algebra
- Matrix product
- Determinants
- Inverse matrices
- Regression using matrix algebra

### Rules

1. You should use no external libraries.

1. Efficiency is essential, but getting it done is even more critical. Get the first version done, and then check for redundant computations.

1. How did I solve it? I implemented one method at a time in the way I described in the questions. The order of the questions is a good predictor of what should be done first and next.

1. The Laplace expansion is computationally expensive.


### Guide to Exams

Programming with math is hard. But it can also be a lot of fun. If you find yourself overwhelmed or frustrated with a question, remember: we're here to help!

## **Have fun!**

***

### 01. **Constructor** (0.5 pts):

You should create a constructor that allows creating the object `Matrix`. 

The user should be able to pass a list or a tuple of numbers and the number of rows of your matrix.

For example, consider the 2 x 3 matrix:

$$
\begin{bmatrix}
1 & 2 & 3 \\
1 & 2 & 4
\end{bmatrix}
$$

The user should be able to create it by doing this:

```python
m = Matrix([1,2,3,1,2,4], 2)
```

You should test whether these elements conform, i. e. if the number of rows and columns is compatible with that of the matrix. 

If numbers conform, you should store the following objects inside the object as class attributes:

```python
print(m.num)
[1,2,3,1,2,4]

print(m.mat)
[[1,2,3],[1,2,4]]

print(m.dim)
[2, 3]
```

Some starting code:

```python
class Matrix:
    """
    Matrix Class
    
    Great to do Linear Algebra!
    """
    # Constructor
    def __init__(self, nums, nrow = None):
        """
        Constructor
        Receives:
            nums: list of numbers
            nrow: number of rows (or None)
        
        Return: 
            Matrix object
        
        """
        pass
```

<!-- grade_id: cell-0cb1a2309a10dee0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Matrix:
    def __init__(self, nums, nrow = None):
        if nrow is None or not isinstance(nrow, int) or nrow <= 0:
            raise ValueError("nrow must be a positive integer")

        if len(nums) % nrow != 0:
            raise ValueError("Cannot evenly divide nums into nrow rows")

        self.num = nums  # list
        self.nrow = nrow  # int
        self.ncol = len(nums) // nrow  # int
        self.mat = self.format()  # list[list]
        self.dim = [self.nrow, self.ncol]  # list
        
    def format(self):
        result = []
        for i in range(0, len(self.num), self.ncol):
            result.append(self.num[i:i + self.ncol])
        return result
```

<!-- grade_id: cell-6c93e184fb56c7d5, points: 0.5, grade: true, solution: false -->
```python
m = Matrix([1, 2, 3, 4, 5, 5, 6, 6], 2)
assert m.mat == [[1, 2, 3, 4], [5, 5, 6, 6]]
assert m.dim == [2, 4]
assert m.num == [1, 2, 3, 4, 5, 5, 6, 6]
try:
    m = Matrix([1, 2, 3, 4, 5], 2)
    assert False
except ValueError:
    assert True
m = Matrix([1, 2, 3, 4], 2)
assert m.mat == [[1, 2], [3, 4]]
assert m.dim == [2, 2]
```

### 02. **\_\_repr\_\_** (0.5 pts):

Implement the **\_\_repr\_\_** for displaying the following:

```python
m
Matrix dimension [2, 3]
```

Some starting code:

```python
class Matrix(Matrix):
    # Representation
    def __repr__(self):
        """
        __repr__:
            Prints the matrix dimensions
        """
        pass
```

<!-- grade_id: cell-5ec42a0a3a7b821f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Matrix:
    def __init__(self, nums, nrow = None):
        if nrow is None or not isinstance(nrow, int) or nrow <= 0:
            raise ValueError("nrow must be a positive integer")

        if len(nums) % nrow != 0:
            raise ValueError("Cannot evenly divide nums into nrow rows")

        self.num = nums
        self.nrow = nrow 
        self.ncol = len(nums) // nrow
        self.mat = self.format()
        self.dim = [self.nrow, self.ncol]
        
    def format(self):
        result = []
        for i in range(0, len(self.num), self.ncol):
            result.append(self.num[i:i + self.ncol])
        return result

    def __repr__(self):
        return f"Matrix dimension {self.dim}"
```

<!-- grade_id: cell-5d7ef63b52342659, points: 0.5, grade: true, solution: false -->
```python
m = Matrix([1, 2, 3, 4, 5, 5, 6, 6], 2)
assert m.__repr__() == 'Matrix dimension [2, 4]'
```

### 03. **`__str__`** (0.5 pts):

Implement an **\_\_str\_\_** representation for your matrix. For example, for the matrix.

$$
\begin{bmatrix}
1 & 2 & 3 \\
1 & 2 & 4
\end{bmatrix}
$$

Should display as:

```python
[1, 2, 3]
[4, 5, 6]
```

Some starting code:

```python
class Matrix(Matrix):
    # Print represetation
    def __str__(self):
        """
        Prints the matrix when using the command print()
        """
        pass
```

<!-- grade_id: cell-1e44ac1a92a9d76b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Matrix:
    def __init__(self, nums, nrow = None):
        if nrow is None or not isinstance(nrow, int) or nrow <= 0:
            raise ValueError("nrow must be a positive integer")

        if len(nums) % nrow != 0:
            raise ValueError("Cannot evenly divide nums into nrow rows")

        self.num = nums
        self.nrow = nrow 
        self.ncol = len(nums) // nrow
        self.mat = self.format()
        self.dim = [self.nrow, self.ncol]
        
    def format(self):
        result = []
        for i in range(0, len(self.num), self.ncol):
            result.append(self.num[i:i + self.ncol])
        return result

    def __repr__(self):
        return f"Matrix dimension {self.dim}"

    def __str__(self):
        result = []
        for i in range(0, len(self.num), self.ncol):
            rows = self.num[i:i + self.ncol]
            result.append(str(rows))
        return '\n'.join(result) + '\n'
```

<!-- grade_id: cell-f32cef9eaeaf8921, points: 0.5, grade: true, solution: false -->
```python
m = Matrix([1, 2, 3, 4, 5, 5, 6, 6], 2)
assert m.__str__() == '[1, 2, 3, 4]\n[5, 5, 6, 6]\n'
```

### 04. **`is_square`** (0.5 pts):

Implement a **is_square** method that returns `True` when the matrix is square.

Example:

```python
m1 = Matrix([1,2,3,1,2,4], 2)
m2 = Matrix([-2,-1,2,2,1,4,-3,3,-1], 3)

print(m1.is_square())
False

print(m2.is_square())
True
```

Some starting code:

```python
class Matrix(Matrix):
    # Test to see if square matrix
    def is_square(self):
        """
        is_square:
            Receives:
                Matrix
                
            Return:
                True if the matrix is square
                False otherwise
        """
        pass
```

<!-- grade_id: cell-9cd3faab4bf82e76, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Matrix:
    def __init__(self, nums, nrow = None):
        if nrow is None or not isinstance(nrow, int) or nrow <= 0:
            raise ValueError("nrow must be a positive integer")

        if len(nums) % nrow != 0:
            raise ValueError("Cannot evenly divide nums into nrow rows")

        self.num = nums
        self.nrow = nrow 
        self.ncol = len(nums) // nrow
        self.mat = self.format()
        self.dim = [self.nrow, self.ncol]
    
    def format(self):
        result = []
        for i in range(0, len(self.num), self.ncol):
            result.append(self.num[i:i + self.ncol])
        return result
        
    def is_square(self):
        if self.nrow == self.ncol:
            return True
        if self.nrow == self.ncol:
            return False
```

<!-- grade_id: cell-6e5889f69e77bf20, points: 0.5, grade: true, solution: false -->
```python
m = Matrix([1, 2, 3, 4, 5, 5, 6, 6], 2)
assert not m.is_square()
```

### 05. **`__eq__`** (0.5 pts):

Implement the **__eq__** method that returns `True` when one matrix is the same as the other.

Example:

```python
m1 = Matrix([1,2,3,1,2,4], 2)
m2 = Matrix([1,2,3,1,2,4], 2)
m3 = Matrix([1,2,3,1,2,4], 3)
m4 = Matrix([2,2,3,1,2,4], 2)

print(m1 == m2)
True

print(m1 == m3)
False

print(m1 == m4)
False

print(m3 == m4)
False

print(m2 == m2)
True
```

Some starting code:

```python
class Matrix(Matrix):
    # Matrix equality
    def __eq__(self, other):
        """
        __eq__:
            Implements a method to check if two matrices are equal to each other
            
            Equal matrices mean:
                1. Same dimension
                2. Each component is the same
            
            Receives:
                Two matrices
                
            Returns
                Boolean
        """
        pass
```

<!-- grade_id: cell-1e3f5a4b4ecd1477, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Matrix:
    def __init__(self, nums, nrow = None):
        if nrow is None or not isinstance(nrow, int) or nrow <= 0:
            raise ValueError("nrow must be a positive integer")

        if len(nums) % nrow != 0:
            raise ValueError("Cannot evenly divide nums into nrow rows")

        self.num = nums
        self.nrow = nrow 
        self.ncol = len(nums) // nrow
        self.mat = self.format()
        self.dim = [self.nrow, self.ncol]
    
    def format(self):
        result = []
        for i in range(0, len(self.num), self.ncol):
            result.append(self.num[i:i + self.ncol])
        return result

    def __eq__(self, other):
        if self.dim != other.dim:
            return False
        elif self.num == other.num:
            return True
        else:
            return False
```

```python
def __eq__(self, other):
    return self.dim == other.dim and self.num == other.num
```

<!-- grade_id: cell-3a9775e1ee4f2d47, points: 0.5, grade: true, solution: false -->
```python
m = Matrix([1, 2, 3, 4, 5, 5, 6, 6], 2)
assert m == m
```

### 06. **`__ne__`** (0.25 pts):

Implement the **`__ne__`** method that returns `True` when the matrices are different from each other.

Example:

```python
m1 = Matrix([1,2,3,1,2,4], 2)
m2 = Matrix([1,2,3,1,2,4], 2)
m3 = Matrix([1,2,3,1,2,4], 3)
m4 = Matrix([2,2,3,1,2,4], 2)

print(m1 != m2)
False

print(m1 != m3)
True

print(m1 != m4)
True

print(m3 != m4)
True

print(m2 != m2)
False
```

Some starting code:

```python
class Matrix(Matrix):
    # Matrix equality
    def __ne__(self, other):
        """
        __ne__:
            Implements a method to check if two matrices are different from each other
            
            Equal matrices mean:
                1. Same dimension
                2. Each component is the same

            Thus, different matrices should have at least one of these components different.
            
            Receives:
                Two matrices
                
            Returns
                Boolean
        """
        pass
```

<!-- grade_id: cell-a233ea61e881c46f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Matrix:
    def __init__(self, nums, nrow = None):
        if nrow is None or not isinstance(nrow, int) or nrow <= 0:
            raise ValueError("nrow must be a positive integer")

        if len(nums) % nrow != 0:
            raise ValueError("Cannot evenly divide nums into nrow rows")

        self.num = nums
        self.nrow = nrow 
        self.ncol = len(nums) // nrow
        self.mat = self.format()
        self.dim = [self.nrow, self.ncol]
    
    def format(self):
        result = []
        for i in range(0, len(self.num), self.ncol):
            result.append(self.num[i:i + self.ncol])
        return result

    def __ne__(self, other):
        return self.dim != other.dim or self.num != other.num
```

<!-- grade_id: cell-ec09e41ac9acc317, points: 0.25, grade: true, solution: false -->
```python
m = Matrix([1, 2, 3, 4, 5, 5, 6, 6], 2)
assert not m != m
```

### 07. Implement a method for scalar multiplication (**\_\_mul\_\_**) (0.5 pts)

In this part, you will implement the method to do scalar multiplications. For the theory regarding scalar multiplication, see: [https://en.wikipedia.org/wiki/Scalar_multiplication].

The multiplication should proceed as follows:

```python
m = Matrix([1,2,3,1,2,4], 2)

print(m)
[1, 2, 3]
[1, 2, 4]

print(m * 2)
[2, 4, 6]
[2, 4, 8]
```

Some starting code:

```python
class Matrix(Matrix):
    # Scalar multiplication (https://en.wikipedia.org/wiki/Scalar_multiplication)
    def __mul__(self, num):
        """
        __mul__:
            Implements scalar multiplication
            
            Receives:
                Matrix
                A float or an int
                
            Returns:
                Matrix
        """
        pass
```

*Hint*: The issue here is with `2 * m`. This should not be defined.

<!-- grade_id: cell-7fd2320cfb1dde9a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Matrix:
    def __init__(self, nums, nrow=None):
        if nrow is None or not isinstance(nrow, int) or nrow <= 0:
            raise ValueError("nrow must be a positive integer")
        if len(nums) % nrow != 0:
            raise ValueError("Cannot evenly divide nums into nrow rows")

        self.num = nums
        self.nrow = nrow
        self.ncol = len(nums) // nrow
        self.mat = self.format()
        self.dim = [self.nrow, self.ncol]

    def format(self):
        result = []
        for i in range(0, len(self.num), self.ncol):
            result.append(self.num[i:i + self.ncol])
        return result

    def __mul__(self, k):
        new_matrix = [i * k for i in self.num]
        return Matrix(new_matrix, self.nrow)

# 没有定义 __eq__，Python 只能判断对象是不是同一个实例，而不是内容是否相等。
    def __eq__(self, other):
        return self.dim == other.dim and self.num == other.num
```

<!-- grade_id: cell-d6166bc8cccbc252, points: 0.5, grade: true, solution: false -->
```python
m = Matrix([1, 2, 3, 4], 2)
assert (m * 2) == Matrix([2, 4, 6, 8], 2)
```

### 08. **\_\_add\_\_** (0.5 pts)

Implement the sum of matrices method (**\_\_add\_\_**). This method ensures that the operations `m1 + m2` is well-defined.

Remember that sum is only defined if the matrices have the same size.

Some starting code:

```python
class Matrix(Matrix):
    # Addition (see https://en.wikipedia.org/wiki/Matrix_addition)
    def __add__(self, other):
        """
        __add__:
            Implements the sum of matrices
            
            Receives:
                Two matrices
                
            Returns
                Subtraction of matrices as a Matrix or error if not conform
        """
        pass
```

<!-- grade_id: cell-75b54c112a457169, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Matrix:
    def __init__(self, nums, nrow=None):
        if nrow is None or not isinstance(nrow, int) or nrow <= 0:
            raise ValueError("nrow must be a positive integer")
        if len(nums) % nrow != 0:
            raise ValueError("Cannot evenly divide nums into nrow rows")

        self.num = nums
        self.nrow = nrow
        self.ncol = len(nums) // nrow
        self.mat = self.format()
        self.dim = [self.nrow, self.ncol]

    def format(self):
        result = []
        for i in range(0, len(self.num), self.ncol):
            result.append(self.num[i:i + self.ncol])
        return result

    def __mul__(self, k):
        new_matrix = [i * k for i in self.num]
        return Matrix(new_matrix, self.nrow)

    def __add__(self, other):
        if self.dim == other.dim:
            new_Matrix = [self.num[i] + other.num[i] for i in range(len(self.num))]
            return Matrix(new_Matrix, self.nrow)
        else:
            ValueError("Not the same dimension")

# 没有定义 __eq__，Python 只能判断对象是不是同一个实例，而不是内容是否相等。
    def __eq__(self, other):
        return self.dim == other.dim and self.num == other.num
```

<!-- grade_id: cell-098f6c6c6a6b759d, points: 0.5, grade: true, solution: false -->
```python
m1 = Matrix([1, 2, 3, 4], 2)
m2 = Matrix([2, 3, -2, -1], 2)
assert (m1 + m2) == Matrix([3, 5, 1, 3], 2)
```

### 09. **\_\_sub\_\_** (0.25 pts)

Implement the subtraction of matrices method (**\_\_sub\_\_**). This method ensures that the operation `m1 - m2` is well-defined.

Remember that subtraction is only defined if the matrices have the same size.

Some starting code:

```python
class Matrix(Matrix):
    # Subtraction
    def __sub__(self, other):
        """
        __sub__:
            Implements the subtraction of two matrices
            
            Receives:
                Two matrices
                
            Returns:
                Subtraction of matrices as a Matrix or error if not conform
        """
        pass
```

<!-- grade_id: cell-29b45ec6de90c7dc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Matrix:
    def __init__(self, nums, nrow=None):
        if nrow is None or not isinstance(nrow, int) or nrow <= 0:
            raise ValueError("nrow must be a positive integer")
        if len(nums) % nrow != 0:
            raise ValueError("Cannot evenly divide nums into nrow rows")

        self.num = nums
        self.nrow = nrow
        self.ncol = len(nums) // nrow
        self.mat = self.format()
        self.dim = [self.nrow, self.ncol]

    def format(self):
        result = []
        for i in range(0, len(self.num), self.ncol):
            result.append(self.num[i:i + self.ncol])
        return result

    def __mul__(self, k):
        new_matrix = [i * k for i in self.num]
        return Matrix(new_matrix, self.nrow)

    def __sub__(self, other):
        if self.dim == other.dim:
            new_Matrix = [self.num[i] - other.num[i] for i in range(len(self.num))]
            return Matrix(new_Matrix, self.nrow)
        else:
            ValueError("Not the same dimension")

# 没有定义 __eq__，Python 只能判断对象是不是同一个实例，而不是内容是否相等。
    def __eq__(self, other):
        return self.dim == other.dim and self.num == other.num
```

<!-- grade_id: cell-b47346a0b04e42b0, points: 0.25, grade: true, solution: false -->
```python
m3 = Matrix([1, 2, 3, 4, 5, 6], 2)
assert (m3 - m3) == Matrix([0, 0, 0, 0, 0, 0], 2)
```

### 10. Implement a method to extract rows (**exrow**) of a matrix (0.5 pts)

For the row extraction, the input is the matrix and a row. It should return a list with the row. One example is:

```python
m = Matrix([1,2,3,1,2,4], 2)

print(m)
[1, 2, 3]
[1, 2, 4]

print(m.exrow(0))
[1, 2, 3]

print(m.exrow(1))
[1, 2, 4]
```

Some starting code:

```python
class Matrix(Matrix):
    # Extract row
    def exrow(self, i):
        """
        exrow:
            Extracts row of a Matrix
            
            Receives:
                Matrix
                Index (int)
                
            Returns:
                List with row or error if index out of bounds
        """
        pass
```

*Note: This method and the next are going to make it easy to do the matrix multiplication.*

<!-- grade_id: cell-f8d53f469aa8fd02, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Matrix:
    def __init__(self, nums, nrow = None):
        if nrow is None or not isinstance(nrow, int) or nrow <= 0:
            raise ValueError("nrow must be a positive integer")

        if len(nums) % nrow != 0:
            raise ValueError("Cannot evenly divide nums into nrow rows")

        self.num = nums
        self.nrow = nrow 
        self.ncol = len(nums) // nrow
        self.mat = self.format()
        self.dim = [self.nrow, self.ncol]
        
    def format(self):
        result = []
        for i in range(0, len(self.num), self.ncol):
            result.append(self.num[i:i + self.ncol])
        return result

    def exrow(self, i):
        if not (0 <= i < self.nrow):
            raise IndexError("Row index out of range")
        return self.mat[i]
```

<!-- grade_id: cell-55140c792c98ad19, points: 0.5, grade: true, solution: false -->
```python
m = Matrix([1, 2, 3, 4, 5, 6], 2)
assert m.exrow(0) == [1, 2, 3]
```

### 13. Implement a method to extract columns (**excol**) of a matrix (0.5 pts)

For the column extraction, the input is the matrix and a column index. It should return a list with the column. One example is:

```
m = Matrix([1,2,3,1,2,4], 2)

print(m)
[1, 2, 3]
[1, 2, 4]

print(m.excol(0))
[1, 1]

print(m.excol(2))
[3, 4]
```

Again, if you choose a column that does not exist, it should throw an error (e.g., 'Error. Invalid column.').

Some starting code:

```python
class Matrix(Matrix):
    # Extract column
    def excol(self, i):
        """
        exrow:
            Extracts column of a Matrix
            
            Receives:
                Matrix
                Index (int)
                
            Returns:
                List with column or error if index out of bounds
        """
        pass
```

<!-- grade_id: cell-36913df426f7da20, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Matrix:
    def __init__(self, nums, nrow = None):
        if nrow is None or not isinstance(nrow, int) or nrow <= 0:
            raise ValueError("nrow must be a positive integer")

        if len(nums) % nrow != 0:
            raise ValueError("Cannot evenly divide nums into nrow rows")

        self.num = nums
        self.nrow = nrow 
        self.ncol = len(nums) // nrow
        self.mat = self.format()
        self.dim = [self.nrow, self.ncol]
        
    def format(self):
        result = []
        for i in range(0, len(self.num), self.ncol):
            result.append(self.num[i:i + self.ncol])
        return result

    def exrow(self, i):
        if not (0 <= i < self.nrow):
            raise IndexError("Row index out of range")
        return self.mat[i]

    def excol(self, i):
        if not (0 <= i < self.ncol):
            raise IndexError("Row index out of range")
        else:
            return [self.mat[c][i] for c in range(self.nrow)]
```

<!-- grade_id: cell-d6b79acc1505af08, points: 0.5, grade: true, solution: false -->
```python
m = Matrix([1, 2, 3, 4, 5, 6], 2)
assert m.excol(0) == [1, 4]
```

### 15. Implement the dot-product (**dot_prod**) as explained below (0.5 pts)

Implement a method for computing the dot-product of two matrices. For two given matrices `m1` and `m2`, the dot-product is a method that receives three arguments (besides the original matrix):

```python
m1.dot_prod(m2, i, j)
```

Or, the matrix 2 (`m2`); the row in matrix one (`i`), and the column in matrix 2 (`j`). It returns a scalar with the value of the dot product. For example:

```python
m1 = Matrix([1, 2, 3, 4, 5, 6], 2)
m2 = Matrix([-1, -2, -3, -4, -5, -6, 1, 2, 3], 3)
```

The dot-product of the column 2 of matrix m1 and the row 1 of matrix m2 is equal to: $1*(-1) + 2*(-4) + 3*1 = -6$

```python
m1.dot_prod(m2, 0, 0)
-6
```

Note that the dot-product:

```python
m2.dot_prod(m1, 1, 0)
```

Is not defined, and should throw an IndexError (e.g., `Error. Indexes do not conform.`). This should be the case when the number of columns in matrix 1 is different than the number of rows in matrix 2.

Some starting code:

```python
class Matrix(Matrix):
    # Dot-product
    dot_prod(self, other, i, j):
        """
        dot_prod:
            Dot-product of row i of m1 and col j of m2
            
            Receives:
                Matrix 1
                Matrix 2
                row index Matrix 1
                column index Matrix 2
                
            Returns:
                float or int
        """
        pass
```

*Note*: If you have no idea what a dot-product is, check [this Wikipedia page](https://en.wikipedia.org/wiki/Dot_product).

<!-- grade_id: cell-00731bada8c53d7c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Matrix:
    def __init__(self, nums, nrow = None):
        if nrow is None or not isinstance(nrow, int) or nrow <= 0:
            raise ValueError("nrow must be a positive integer")

        if len(nums) % nrow != 0:
            raise ValueError("Cannot evenly divide nums into nrow rows")

        self.num = nums
        self.nrow = nrow 
        self.ncol = len(nums) // nrow
        self.mat = self.format()
        self.dim = [self.nrow, self.ncol]
        
    def format(self):
        result = []
        for i in range(0, len(self.num), self.ncol):
            result.append(self.num[i:i + self.ncol])
        return result

    def exrow(self, i):
        if not (0 <= i < self.nrow):
            raise IndexError("Row index out of range")
        return self.mat[i]

    def excol(self, i):
        if not (0 <= i < self.ncol):
            raise IndexError("Row index out of range")
        else:
            return [self.mat[c][i] for c in range(self.nrow)]

    def dot_prod(self, other, i, j):
        if self.ncol != other.nrow:
            raise IndexError("Error. Indexes do not conform.")
        else:
            row = self.exrow(i)
            col = other.excol(j)
            dot_prod = sum(row[k] * col[k] for k in range(self.ncol)) 
            return dot_prod
```

<!-- grade_id: cell-e363f75abdcbea5c, points: 0.5, grade: true, solution: false -->
```python
m1 = Matrix([1, 2, 3, 4, 5, 6], 2)
m2 = Matrix([-1, -2, -3, -4, -5, -6, 1, 2, 3], 3)
assert m1.dot_prod(m2, 0, 0) == -6
```

```python
# zip
```

### 16. Implement the matrix multiplication method (**`__pow__`**) (1 pt)

For learning how matrix multiplication works, check: https://en.wikipedia.org/wiki/Matrix_multiplication

In this method, you will implement the following operation:

```python
m1 = Matrix([1,2,3,4,5,6,7,8,9], 3)
m2 = Matrix([1,2,3,1,2,4], 3)

print(m1)
[1, 2, 3]
[4, 5, 6]
[7, 8, 9]

print(m2)
[1, 2]
[3, 1]
[2, 4]

print(m1 ** m2)
[13, 16]
[31, 37]
[49, 58]
```

*Note 1*: If you try `m2 ** m1`, it should throw an error (e.g. 'Error. Matrices do not conform for multiplication.').

Some starting code:

```python
class Matrix(Matrix):
    # Matrix multiplication (see https://en.wikipedia.org/wiki/Matrix_multiplication)
    def __pow__(self, other):
        """
        __pow__:
            Implements the product of two matrices
            
            Receives:
                Two Matrices
                
            Returns:
                Matrix or error if not conform
        """
        pass
```

**Hint:** You can use the **excol** and **exrow** methods to extract the rows and columns. Then, check the definition of a dot product: https://en.wikipedia.org/wiki/Dot_product. This should be exactly what you should do with each row and column extracted.

<!-- grade_id: cell-a723f97babc49461, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Matrix:
    def __init__(self, nums, nrow = None):
        if nrow is None or not isinstance(nrow, int) or nrow <= 0:
            raise ValueError("nrow must be a positive integer")

        if len(nums) % nrow != 0:
            raise ValueError("Cannot evenly divide nums into nrow rows")

        self.num = nums
        self.nrow = nrow 
        self.ncol = len(nums) // nrow
        self.mat = self.format()
        self.dim = [self.nrow, self.ncol]
        
    def format(self):
        result = []
        for i in range(0, len(self.num), self.ncol):
            result.append(self.num[i:i + self.ncol])
        return result

    def exrow(self, i):
        if not (0 <= i < self.nrow):
            raise IndexError("Row index out of range")
        return self.mat[i]

    def excol(self, i):
        if not (0 <= i < self.ncol):
            raise IndexError("Row index out of range")
        else:
            return [self.mat[c][i] for c in range(self.nrow)]

    def dot_prod(self, other, i, j):
        if self.ncol != other.nrow:
            raise IndexError("Error. Indexes do not conform.")
        else:
            row = self.exrow(i)
            col = other.excol(j)
            dot_prod = sum(row[k] * col[k] for k in range(self.ncol)) 
            return dot_prod

    def __pow__(self, other): 
        if self.ncol != other.nrow:
            raise IndexError("Error. Indexes do not conform.")
        elements = []
        for i in range(self.nrow):
            for j in range(other.ncol):
                elements.append(self.dot_prod(other, i, j))
        return Matrix(elements, self.nrow)

# 没有定义 __eq__，Python 只能判断对象是不是同一个实例，而不是内容是否相等。
    def __eq__(self, other):
        return self.dim == other.dim and self.num == other.num
```

<!-- grade_id: cell-665a29c0ab0cb0f1, points: 1, grade: true, solution: false -->
```python
m1 = Matrix([1,0,0,0,1,0,0,0,1], 3)
m2 = Matrix([1,2,3,4,5,6,7,8,9], 3)
assert (m1 ** m1) == m1
assert (m1 ** m2) == m2
```

### 18. Implement the submatrix method for computing submatrices (**subm**) (0.5 pts)

To compute the minor of a matrix, you have to extract a submatrix. For example, if the matrix is:

```python
m = Matrix([1,2,3,4,5,6,7,8,9], 3)

print(m)
[1, 2, 3]
[4, 5, 6]
[7, 8, 9]
```

The minor $M_{2,2}$ is the determinant of the submatrix:

```python
print(m.subm(1,1))
[1, 3]
[7, 9]
```

I.e., to compute the submatrix, you should delete the second row and column:

```python
[1, -, 3]
[-, -, -]
[7, -, 9]
```

And return the remaining numbers as a matrix.

Some starting code:

```python
class Matrix(Matrix):
    # Extract submatrix (see https://en.wikipedia.org/wiki/Minor_(linear_algebra))
    def subm(self, row, col):
        """
        subm:
            Extracts a submatrix (for minor and cofactor computations)
            
            Receives:
                Matrix
                row (int)
                col (int)
                
            Returns:
                Matrix
        """
        pass
```

***Hint*** Are there any methods already implemented that would help you with this?

<!-- grade_id: cell-6a3c9cd272f2d103, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Matrix:
    def __init__(self, nums, nrow = None):
        if nrow is None or not isinstance(nrow, int) or nrow <= 0:
            raise ValueError("nrow must be a positive integer")

        if len(nums) % nrow != 0:
            raise ValueError("Cannot evenly divide nums into nrow rows")

        self.num = nums
        self.nrow = nrow 
        self.ncol = len(nums) // nrow
        self.mat = self.format()
        self.dim = [self.nrow, self.ncol]
        
    def format(self):
        result = []
        for i in range(0, len(self.num), self.ncol):
            result.append(self.num[i:i + self.ncol])
        return result

    def exrow(self, i):
        if not (0 <= i < self.nrow):
            raise IndexError("Row index out of range")
        return self.mat[i]

    def excol(self, i):
        if not (0 <= i < self.ncol):
            raise IndexError("Row index out of range")
        else:
            return [self.mat[c][i] for c in range(self.nrow)]

    def subm(self, row, col):
        if not (0<=row<self.nrow and 0<=row<self.ncol):
            raise IndexError("Index out of range")
        result = []
        for i in range(self.nrow):
            for j in range(self.ncol):
                if i != row and j != col:
                    result.append(self.mat[i][j])
        return Matrix(result, self.nrow-1)

# 没有定义 __eq__，Python 只能判断对象是不是同一个实例，而不是内容是否相等。
    def __eq__(self, other):
        return self.dim == other.dim and self.num == other.num
```

<!-- grade_id: cell-bd09d5a3068e422a, points: 0.5, grade: true, solution: false -->
```python
m = Matrix([1,2,3,4,5,6,7,8,9], 3)
assert m.subm(1, 1) == Matrix([1, 3, 7, 9], 2)
assert m.subm(1, 1).subm(0,0) == Matrix([9], 1)
```

### 19. Implement the determinant method (**det**) (1 pt)

First, if the user tries to compute the determinant in a non-square matrix, the code should throw an error (e.g. 'Error. Matrix is not square.'). Else, it should start the computations.

If you don't know what a recursive function is, please read the following entry: https://pythonnumericalmethods.berkeley.edu/notebooks/chapter06.01-Recursive-Functions.html

If you don't remember much about how to compute the determinant, check the cofactor expansion method here: https://en.wikipedia.org/wiki/Minor_(linear_algebra)

We will compute the determinant by using a method called cofactor expansion. It consists of computing the determinant recursively, using the relevant cofactors.

The way I solved was:

```
...some code before...
if --matrix_is_1x1--:
    return --we-know-the-determinant-here--
else:
    for --iteration-in-expansion--:
        --computations-and-recursion--
...some code after...
```

One code that could guide you is in here: https://en.wikipedia.org/wiki/Laplace_expansion

To use as examples:

```python
m = Matrix([-2,-1,2,2,1,4,-3,3,-1], 3)

print(m)
[-2, -1, 2]
[2, 1, 4]
[-3, 3, -1]

m.det()
54
```

And:

```python
m = Matrix([1,2,3,4,5,6,7,8,9], 3)

print(m)
[1, 2, 3]
[4, 5, 6]
[7, 8, 9]

m.det()
0
```

Some starting code:

```python
class Matrix(Matrix):
    # Determinant (see https://en.wikipedia.org/wiki/Determinant)
    def det(self):
        """
        det:
            Compute the determinant of a matrix by cofactor expansion
            
            Receives:
                Matrix
                
            Returns:
                float, int or an error if matrix is not square.
        """
        pass
```

<!-- grade_id: cell-91fc434d5c3c16fb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Matrix:
    def __init__(self, nums, nrow = None):
        if nrow is None or not isinstance(nrow, int) or nrow <= 0:
            raise ValueError("nrow must be a positive integer")

        if len(nums) % nrow != 0:
            raise ValueError("Cannot evenly divide nums into nrow rows")

        self.num = nums
        self.nrow = nrow 
        self.ncol = len(nums) // nrow
        self.mat = self.format()
        self.dim = [self.nrow, self.ncol]
        
    def format(self):
        result = []
        for i in range(0, len(self.num), self.ncol):
            result.append(self.num[i:i + self.ncol])
        return result

    def exrow(self, i):
        if not (0 <= i < self.nrow):
            raise IndexError("Row index out of range")
        return self.mat[i]

    def excol(self, i):
        if not (0 <= i < self.ncol):
            raise IndexError("Row index out of range")
        else:
            return [self.mat[c][i] for c in range(self.nrow)]

    def subm(self, row, col):
        if not (0<=row<self.nrow and 0<=row<self.ncol):
            raise IndexError("Index out of range")
        result = []
        for i in range(self.nrow):
            for j in range(self.ncol):
                if i != row and j != col:
                    result.append(self.mat[i][j])
        return Matrix(result, self.nrow-1)

    def det(self):
        if self.nrow != self.ncol:
            raise IndexError("Error. Matrix is not square.")
        if self.nrow == self.ncol == 1:
            return int(self.num)
        if self.nrow == self.ncol == 2:
            det_2_2 = self.mat[0][0] * self.mat[1][1] - self.mat[1][0] * self.mat[0][1]
            return det_2_2
        determinant = 0
        for col in range(self.ncol):
            cofactor = (-1)**col * self.mat[0][col] 
            sub = self.subm(0,col)
            determinant += cofactor * sub.det()
            return determinant

# 没有定义 __eq__，Python 只能判断对象是不是同一个实例，而不是内容是否相等。
    def __eq__(self, other):
        return self.dim == other.dim and self.num == other.num
```

<!-- grade_id: cell-c18eb171dee8d0c0, points: 1, grade: true, solution: false -->
```python
m = Matrix([1, 2, 3, 4], 2)
assert m.det() == -2
m = Matrix([1, 2, 1, 2], 2)
assert m.det() == 0
```

### 20. Cofactor matrix calculator (**cofm**) (0.5 pts)

To compute the cofactors, you should use the determinants and submatrices above. Example of cofactor computations:

```python
m = Matrix([1,2,3,4,5,6,7,8,9], 3)

print(m)
[1, 2, 3]
[4, 5, 6]
[7, 8, 9]

print(m.cofm())
[-3, 6, -3]
[6, -12, 6]
[-3, 6, -3]
```

Or, for a non-singular matrix:

```python
m = Matrix([1,2,3,4,5,6,7,8,9], 3)

print(m)
[-2, -1, 2]
[2, 1, 4]
[-3, 3, -1]

print(m.cofm())
[-13, -10, 9]
[5, 8, 9]
[-6, 12, 0]
```

Some starting code:

```python
class Matrix(Matrix):
    # Cofactor matrix (see https://en.wikipedia.org/wiki/Minor_(linear_algebra))
    def cofm(self):
        """
        cofm:
            Computes the Cofactor Matrix
            
            Receives:
                Matrix
                
            Returns:
                Matrix or error if matrix not square.
        """
        pass
```

<!-- grade_id: cell-0114cff7843d25e4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Matrix:
    def __init__(self, nums, nrow = None):
        if nrow is None or not isinstance(nrow, int) or nrow <= 0:
            raise ValueError("nrow must be a positive integer")

        if len(nums) % nrow != 0:
            raise ValueError("Cannot evenly divide nums into nrow rows")

        self.num = nums
        self.nrow = nrow 
        self.ncol = len(nums)// self.nrow
        self.mat = self.format()
        self.dim = [self.nrow, self.ncol]
        
    def format(self):
        result = []
        for i in range(0, len(self.num), self.ncol):
            result.append(self.num[i:i + self.ncol])
        return result

    def exrow(self, i):
        if not (0 <= i < self.nrow):
            raise IndexError("Row index out of range")
        return self.mat[i]

    def excol(self, i):
        if not (0 <= i < self.ncol):
            raise IndexError("Row index out of range")
        else:
            return [self.mat[c][i] for c in range(self.nrow)]

    def subm(self, row, col):
        if not (0<=row<self.nrow and 0<=row<self.ncol):
            raise IndexError("Index out of range")
        result = []
        for i in range(self.nrow):
            for j in range(self.ncol):
                if i != row and j != col:
                    result.append(self.mat[i][j])
        return Matrix(result, self.nrow-1)

    def det(self):
        if self.nrow != self.ncol:
            raise IndexError("Error. Matrix is not square.")
        if self.nrow == self.ncol == 1:
            return self.mat[0][0]
        if self.nrow == self.ncol == 2:
            det_2_2 = self.mat[0][0] * self.mat[1][1] - self.mat[1][0] * self.mat[0][1]
            return det_2_2
        determinant = 0
        for col in range(self.ncol):
            cofactor = (-1)**col * self.mat[0][col] 
            sub = self.subm(0,col)
            determinant += cofactor * sub.det()
            return determinant

    def cofm(self):
        if self.nrow != self.ncol:
            raise IndexError("Error. Matrix is not square.")
        cofactors = []
        for i in range(self.nrow):
            for j in range(self.ncol):
                sub = self.subm(i, j)    
                cof = (-1)**(i+j) * sub.det()
                cofactors.append(cof)
        return Matrix(cofactors, self.nrow)
            
# 没有定义 __eq__，Python 只能判断对象是不是同一个实例，而不是内容是否相等。
    def __eq__(self, other):
        return self.dim == other.dim and self.num == other.num
```

<!-- grade_id: cell-64a123905b510ba5, points: 0.5, grade: true, solution: false -->
```python
m = Matrix([1, 2, 3, 4, 5, 6, 7, 8, 9], 3)
assert m.cofm().dim == [3, 3]
m = Matrix([1, 9, 3, 2, 5, 4, 3, 7, 8], 3)
assert m.cofm().num == [12, -4, -1, -51, -1, 20, 21, 2, -13]
```

### 21. Transpose calculator (**t**) (0.25 pts) 

You need to implement a method to compute the transpose. If you are not sure what a transpose matrix is, please read: https://en.wikipedia.org/wiki/Transpose

For example:

```python
m = Matrix([1,2,3,4,5,6,7,8,9], 3)

print(m)
[-2, -1, 2]
[2, 1, 4]
[-3, 3, -1]

print(m.cofm())
[-13, -10, 9]
[5, 8, 9]
[-6, 12, 0]

print(m.cofm().t())
[-13, 5, -6]
[-10, 8, 12]
[9, 9, 0]
```

*Note*: The transpose of the cofactor matrix is called the *adjoint matrix*.

Some starting code:

```python
class Matrix(Matrix):
    # Transpose
    def t(self):
        """
        t:
            Computes the transpose of a matrix
            
            Receives:
                Matrix
                
            Returns:
                Matrix
        """
        pass
```

<!-- grade_id: cell-94da55352601bf00, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Matrix:
    def __init__(self, nums, nrow = None):
        if nrow is None or not isinstance(nrow, int) or nrow <= 0:
            raise ValueError("nrow must be a positive integer")

        if len(nums) % nrow != 0:
            raise ValueError("Cannot evenly divide nums into nrow rows")

        self.num = nums
        self.nrow = nrow 
        self.ncol = len(nums)// self.nrow
        self.mat = self.format()
        self.dim = [self.nrow, self.ncol]
        
    def format(self):
        result = []
        for i in range(0, len(self.num), self.ncol):
            result.append(self.num[i:i + self.ncol])
        return result

    def exrow(self, i):
        if not (0 <= i < self.nrow):
            raise IndexError("Row index out of range")
        return self.mat[i]

    def excol(self, i):
        if not (0 <= i < self.ncol):
            raise IndexError("Row index out of range")
        else:
            return [self.mat[c][i] for c in range(self.nrow)]

    def subm(self, row, col):
        if not (0<=row<self.nrow and 0<=row<self.ncol):
            raise IndexError("Index out of range")
        result = []
        for i in range(self.nrow):
            for j in range(self.ncol):
                if i != row and j != col:
                    result.append(self.mat[i][j])
        return Matrix(result, self.nrow-1)

    def det(self):
        if self.nrow != self.ncol:
            raise IndexError("Error. Matrix is not square.")
        if self.nrow == self.ncol == 1:
            return self.mat[0][0]
        if self.nrow == self.ncol == 2:
            det_2_2 = self.mat[0][0] * self.mat[1][1] - self.mat[1][0] * self.mat[0][1]
            return det_2_2
        determinant = 0
        for col in range(self.ncol):
            cofactor = (-1)**col * self.mat[0][col] 
            sub = self.subm(0,col)
            determinant += cofactor * sub.det()
            return determinant

    def cofm(self):
        if self.nrow != self.ncol:
            raise IndexError("Error. Matrix is not square.")
        cofactors = []
        for i in range(self.nrow):
            for j in range(self.ncol):
                sub = self.subm(i, j)    
                cof = (-1)**(i+j) * sub.det()
                cofactors.append(cof)
        return Matrix(cofactors, self.nrow)

    def t(self):
        result = []
        for j in range(self.ncol):
            for i in range(self.nrow):
                result.append(self.mat[i][j])
        return Matrix(result, self.ncol)
            
# 没有定义 __eq__，Python 只能判断对象是不是同一个实例，而不是内容是否相等。
    def __eq__(self, other):
        return self.dim == other.dim and self.num == other.num
```

<!-- grade_id: cell-5decdc4ad82f1cd1, points: 0.25, grade: true, solution: false -->
```python
m = Matrix([1, 2, 3, 4, 5, 6], 2)
assert m.t().dim == [m.dim[1], m.dim[0]]
m = Matrix([1, 2, 3, 4], 2)
assert m.t().det() == m.det()
```

### 22. Compute the inverse matrix calculator (**inv**) (0.25 pts):

If you did all the steps before, the inverse matrix is going to be straightforward:

$$ A^{-1} = \dfrac{1}{det(A)}C^t $$

Where $C$ is the cofactor matrix. For example:

```
m = Matrix([1,2,3,4,5,6,7,8,9], 3)

print(m)
[-2, -1, 2]
|2, 1, 4|
[-3, 3, -1]

print(m.inv())
[-0.24074074074074073, 0.09259259259259259, -0.1111111111111111]
|-0.18518518518518517, 0.14814814814814814, 0.2222222222222222|
[0.16666666666666666, 0.16666666666666666, 0.0]
```

Some starting code:

```python
class Matrix(Matrix):
    # Transpose
    def inv(self):
        """
        inv:
            Computes Inverse Matrix using cofactor matrix
            
            Receives:
                Matrix
                
            Returns:
                Matrix or error if matrix is singular or not square.
        """
        pass
```

<!-- grade_id: cell-dddd6feb57274efd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# YOUR CODE HERE
class Matrix:
    def __init__(self, nums, nrow = None):
        if nrow is None or not isinstance(nrow, int) or nrow <= 0:
            raise ValueError("nrow must be a positive integer")

        if len(nums) % nrow != 0:
            raise ValueError("Cannot evenly divide nums into nrow rows")

        self.num = nums
        self.nrow = nrow 
        self.ncol = len(nums)// self.nrow
        self.mat = self.format()
        self.dim = [self.nrow, self.ncol]
        
    def format(self):
        result = []
        for i in range(0, len(self.num), self.ncol):
            result.append(self.num[i:i + self.ncol])
        return result

    def exrow(self, i):
        if not (0 <= i < self.nrow):
            raise IndexError("Row index out of range")
        return self.mat[i]

    def excol(self, i):
        if not (0 <= i < self.ncol):
            raise IndexError("Row index out of range")
        else:
            return [self.mat[c][i] for c in range(self.nrow)]

    def __mul__(self, k):
        new_matrix = [i * k for i in self.num]
        return Matrix(new_matrix, self.nrow)
        
    def subm(self, row, col):
        if not (0<=row<self.nrow and 0<=row<self.ncol):
            raise IndexError("Index out of range")
        result = []
        for i in range(self.nrow):
            for j in range(self.ncol):
                if i != row and j != col:
                    result.append(self.mat[i][j])
        return Matrix(result, self.nrow-1)

    def det(self):
        if self.nrow != self.ncol:
            raise IndexError("Error. Matrix is not square.")
        if self.nrow == self.ncol == 1:
            return self.mat[0][0]
        if self.nrow == self.ncol == 2:
            det_2_2 = self.mat[0][0] * self.mat[1][1] - self.mat[1][0] * self.mat[0][1]
            return det_2_2
        determinant = 0
        for col in range(self.ncol):
            cofactor = (-1)**col * self.mat[0][col] 
            sub = self.subm(0,col)
            determinant += cofactor * sub.det()
            return determinant

    def cofm(self):
        if self.nrow != self.ncol:
            raise IndexError("Error. Matrix is not square.")
        cofactors = []
        for i in range(self.nrow):
            for j in range(self.ncol):
                sub = self.subm(i, j)    
                cof = (-1)**(i+j) * sub.det()
                cofactors.append(cof)
        return Matrix(cofactors, self.nrow)

    def t(self):
        result = []
        for j in range(self.ncol):
            for i in range(self.nrow):
                result.append(self.mat[i][j])
        return Matrix(result, self.ncol)

    def inv(self):
        if self.det() == 0:
            raise ZeroDivisionError("Matrix is not invertible (det = 0)")
        return self.cofm().t() * (1 / self.det())

# 没有定义 __eq__，Python 只能判断对象是不是同一个实例，而不是内容是否相等。
    def __eq__(self, other):
        return self.dim == other.dim and self.num == other.num
```

<!-- grade_id: cell-c19f7c84ffc95ff2, points: 0.25, grade: true, solution: false -->
```python
m = Matrix([1, 2, 3, 4], 2)
assert m.inv().dim == [2, 2]
assert m.inv().num == [-2.0, 1.0, 1.5, -0.5]
```

### 23. Compute the regression coefficients for the following matrix (1 pt):

The formula for the regression is:

$$ \beta = (X^TX)^{-1}X^Ty $$

Where $X^T$ is the transpose of $X$. If we have the following data matrix **X**:

```
X = Matrix([1, -4, 0, 1, -3, 0, 1, -2, 0, 1, -1, 0, 1, 0, 0, 1, 1, 1, 1, 2, 1, 1, 3, 1, 1, 4, 1, 1, 5, 1], 10)
```

And the data **y**:

```
y = Matrix([-3.56518269407931, -2.96460021969053, -1.48659413971178, 0.700322671620325, 0.902347734124384, 1.29721825127923, 1.92302031156774, 3.88376671356276, 2.23958110933972, 6.5610333044646], 10)
```

This represents the formula:

$$ y = X\beta + \varepsilon $$

And for each individual coefficient:

$$ y_i = \beta_0 + \beta_1 x_{1i} + \beta_2 x_{2i} + \varepsilon_i $$

Computing the formula on these data, you should find a vector of coefficients close to:

```
[1.0616761356387148]
[1.1722087325930488]
[-1.3973783953750463]
```

Which represents:

$$ y_i = 1.06 + 1.17 x_{1i} - 1.40 x_{2i} + \varepsilon_i $$

*Congrats! You just created a Python package to compute regressions!*

```python
X = Matrix([1, -4, 0, 1, -3, 0, 1, -2, 0, 1, -1, 0, 1, 0, 0, 1, 1, 1, 1, 2, 1, 1, 3, 1, 1, 4, 1, 1, 5, 1], 10)
y = Matrix([-3.56518269407931, -2.96460021969053, -1.48659413971178, 0.700322671620325, 0.902347734124384, 1.29721825127923, 1.92302031156774, 3.88376671356276, 2.23958110933972, 6.5610333044646], 10)
print((X.t()**X).inv()**X.t()**y)
```

Output:
```
TypeError                                 Traceback (most recent call last)
Cell In[102], line 3
      1 X = Matrix([1, -4, 0, 1, -3, 0, 1, -2, 0, 1, -1, 0, 1, 0, 0, 1, 1, 1, 1, 2, 1, 1, 3, 1, 1, 4, 1, 1, 5, 1], 10)
      2 y = Matrix([-3.56518269407931, -2.96460021969053, -1.48659413971178, 0.700322671620325, 0.902347734124384, 1.29721825127923, 1.92302031156774, 3.88376671356276, 2.23958110933972, 6.5610333044646], 10)
----> 3 print((X.t()**X).inv()**X.t()**y)

TypeError: unsupported operand type(s) for ** or pow(): 'Matrix' and 'Matrix'
```

<!-- grade_id: cell-d23977294e14b23e, points: 1, grade: true, solution: false -->
```python
assert ((X.t()**X).inv()**X.t()**y).num[0] >= 1.06
assert ((X.t()**X).inv()**X.t()**y).num[1] >= 1.17
assert ((X.t()**X).inv()**X.t()**y).num[2] <= -1.39
```

Output:
```
TypeError                                 Traceback (most recent call last)
Cell In[103], line 1
----> 1 assert ((X.t()**X).inv()**X.t()**y).num[0] >= 1.06
      2 assert ((X.t()**X).inv()**X.t()**y).num[1] >= 1.17
      3 assert ((X.t()**X).inv()**X.t()**y).num[2] <= -1.39

TypeError: unsupported operand type(s) for ** or pow(): 'Matrix' and 'Matrix'
```

## Submit!

Once you've completed all the cells above (saving regularly):

- Click "Validate". This will run a check to determine whether you've passed all visible tests. 
- Once you've validated the assignment, you should now have an option to "Submit" the assignment (next to where the assignment is stored in your directory). Click this.
- This will now appear under your "Submitted Assignments" section.

If you have any trouble accessing or submitting the assignment, please check in with your TA!